# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

```python
# imports
import pandas as pd
import numpy as np
```

```python
df = pd.read_csv('assets/NISPUF17.csv')
```

```python
print(df.shape)
df.head()
```

Output:
```
(28465, 454)
   Unnamed: 0  SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  \
0           1   128521     12852     2         NaN  235.916956     1031  2017   
1           2    10741      1074     2         NaN  957.353840     1068  2017   
2           3   220011     22001     2         NaN  189.611299     1050  2017   
3           4    86131      8613     1  675.430817  333.447418     1040  2017   
4           5   227141     22714     1  482.617748  278.768063     1008  2017   

   AGECPOXR  HAD_CPOX  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  XVRCTY6  \
0       NaN         2  ...                        NaN      NaN      NaN   
1       NaN         2  ...                        NaN      NaN      NaN   
2       NaN         2  ...                        NaN      NaN      NaN   
3       NaN         2  ...                        NaN      NaN      NaN   
4       NaN         2  ...                        NaN      NaN      NaN   

   XVRCTY7  XVRCTY8  XVRCTY9  INS_STAT2_I  INS_BR
```

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


```python
# check values
df['EDUC1'].value_counts()
```

Output:
```
4    13656
3     6999
2     4906
1     2904
Name: EDUC1, dtype: int64
```

```python
#calculate ratio
df['EDUC1'].value_counts()/df['EDUC1'].shape[0]
```

Output:
```
4    0.479747
3    0.245881
2    0.172352
1    0.102020
Name: EDUC1, dtype: float64
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
# according to PDF, education information can be found in column EDUC1

def proportion_of_education():
    # YOUR CODE HERE
    edu = df['EDUC1'].value_counts()/df['EDUC1'].shape[0]
    levels = ['less than high school', 'high school', 'more than high school but not college', 'college']
    dic = {}
    i = 1
    for x in levels:
        dic[x] = edu[i] 
        i += 1
    
    return dic
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
# check values // I'm interested only in 1 and 2
df['CBF_01'].value_counts()
```

Output:
```
1     24261
2      4085
77      110
99        9
Name: CBF_01, dtype: int64
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
# according to PDF, breastmilk information can be found in column CBF_01 and seasonal influenza vaccine in column P_NUMFLU

def average_influenza_doses():
    # YOUR CODE HERE
    
    df1 = df[df['CBF_01'] == 1] #where breastfeeded
    df2 = df[df['CBF_01'] == 2] #where non-breastfeeded
    
    return np.mean(df1['P_NUMFLU']), np.mean(df2['P_NUMFLU'])
    #raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
# Sex information can be found in 'SEX' column, Chickenpox in 'HAD_CPOX' and vaccination in 'P_NUMVRC'

def chickenpox_by_sex():
    # YOUR CODE HERE
    df_yes_male = df[(df['P_NUMVRC'] >=1) & (df['HAD_CPOX'] == 1) & (df['SEX'] == 1)] #males who were vaccinated and had chickenpox
    df_yes_famale = df[(df['P_NUMVRC'] >=1) & (df['HAD_CPOX'] == 1) & (df['SEX'] == 2)] #females who were vaccinated and had chickenpox
    df_no_male = df[(df['P_NUMVRC'] >=1) & (df['HAD_CPOX'] == 2) & (df['SEX'] == 1)] #males who were vaccinated and hadn't chickenpox
    df_no_famale = df[(df['P_NUMVRC'] >=1) & (df['HAD_CPOX'] == 2) & (df['SEX'] == 2)] #females who were vaccinated and hadn't chickenpox
    
    ratio_male = df_yes_male.shape[0]/df_no_male.shape[0]
    ratio_female = df_yes_famale.shape[0]/df_no_famale.shape[0]
    ratio = [ratio_male, ratio_female]
    
    sex = ['male', 'female']
    dic = {}
    i = 0
    for x in sex:
        dic[x] = ratio[i]
        i+=1
        
    return dic  
    #raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. If the `had_chickenpox_column` is either `1` (for yes) or `2` for no, and that the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine, then a positive correlation (e.g. `corr > 0`) would mean that an increase in `had_chickenpox_column` (which means more no's) would mean an increase in the `num_chickenpox_vaccine_column` (which means more doses of vaccine). If `corr < 0` then there is a negative correlation, indicating that having had chickenpox is related to an increase in the number of vaccine doses. Also, `pval` refers to the probability the relationship observed is significant. In this case `pval` should be very very small (will end in `e-18` indicating a very small number), which means the result unlikely to be by chance.

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

```python
df['P_NUMVRC'].value_counts()
```

Output:
```
1.0    13781
0.0     1338
2.0      213
3.0        1
Name: P_NUMVRC, dtype: int64
```

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
# As it has been already mentioned, Chickenpox information can be found in 'HAD_CPOX' column and vaccination in 'P_NUMVRC'

def corr_chickenpox():
    import scipy.stats as stats
    
    df_new = df[(df['P_NUMVRC'] >=0) & (df['HAD_CPOX'] <= 2)]
    
    no_or_yes = df_new['HAD_CPOX']
    
    vaccine = df_new['P_NUMVRC']
    
    
    # this is just an example dataframe
    df_final=pd.DataFrame({"had_chickenpox_column":no_or_yes,
                   "num_chickenpox_vaccine_column":vaccine})

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df_final["had_chickenpox_column"],df_final["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    return corr

    #raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python

```